---
output: 
  learnr::tutorial:
    progressive: true
runtime: shiny_prerendered
title: "Appendix B" 
subtitle: |
  | Fundamentals of Probability
author: 
  name: "Created by [Jim Bang](http://www.github.com/bangecon)"
  email: BangJamesT@sau.edu
  affiliation: "[St. Ambrose University](http://www.sau.edu)"
date: "`r format(Sys.Date(), '%d %b %Y')`"
always_allow_html: yes
description: |
  This tutorial practices probability and probability distributions.
---

<style type="text/css">
h1{font-size: 24pt}
h2{font-size: 20pt}
h3{font-size: 18pt}
h4,h5,h6{font-size: 16pt}
body{font-size: 16pt}
</style>

```{r setup, include=FALSE}
library(learnr)
library(gradethis)
gradethis_setup()
tutorial_options(exercise.reveal_solution = FALSE)
knitr::opts_chunk$set(echo = FALSE)
```

## Appendix B, #5

Prior to jury selection for O.J Simpson's murder trial in 1995, a poll found that about 20% of the adult population believed Simpson was innocent (after much of the physical evidence in the case had been revealed to the public). Assume that the 12 jurors were selected randomly and independently from the population (although this turned out not to be true). 

1. Calculate the probability that the jury had at least one member who believed in Simpson's innocence prior to jury selection. 
2. Calculate the probability that the jury had at least two members who believed in Simpson's innocence.

```{r OJ, exercise=TRUE}

```

```{r OJ-solution}
1 - pbinom(0, 12, 0.2)
1 - pbinom(1, 12, 0.2)
```

```{r OJ-hint}
# The number of jurors out of 12 who believe Simpson is innocent follows a binomial distribution.
# P{X >= 1) = 1 - P(X < 1); P(X >= 2) = 1 - P(X < 2)
```

```{r OJ-check}
grade_this({
  # custom checking code appears here
  if (identical(.result, .solution)) {
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

## Exercise

Let X denote the prison sentence, in years (up to a hypothetical three-year maximum sentence), for people convicted of auto theft in a particular state in the United States. Suppose that the pdf of X is given by: 

$$ f(x) = \begin{cases} \begin{align} &(1/9)x^2 &\text{ if } 0 < x < 3 \\ 
  &0 &\text{ otherwise} \end{align} \end{cases}$$

Find the expected value and standard deviation of the prison sentence. 

```{r prison, exercise=TRUE}

```

```{r prison-solution}
xf.x <- function(x) {(1/9)*x^3}
x2f.x <- function(x) {(1/9)*x^4}
integrate(xf.x, 0, 3)$value
sqrt(integrate(x2f.x, 0, 3)$value - integrate(xf.x, 0, 3)$value^2)
```

```{r prison-hint}
# E(x) = integral over sample space of x*f(x), so you will first need to define xf(x). 
# You will also need to define x^2*f(x) to get the variance. 
# Use "integrate(function, lower bound, upper bound)" to get the integral
# V(X) = E(X-mu)^2 = E(X^2) - E(X)^2
# To take the square root of the value of the integral instead of the full result (which includes the error tolerance), extract the $value element from it.
```

```{r prison-check}
grade_this({
  # custom checking code appears here
  if (identical(.result, .solution)) {
    pass(random_praise())
  }
  fail(random_encouragement())
})
```

