---
output: slidy_presentation
runtime: shiny_prerendered
title: "Chapter 7 \n\n Regression with Qualitative Information"
author: "James Bang"
description: >
  This tutorial introduces binary regressors (dummy variables).
---

```{r setup, include=FALSE}
library(learnr)
library(gradethis)
library(wooldridge)
wage1 <- wooldridge::wage1
wage.lm7 <- lm(wage ~ educ + exper + I(exper^2) + tenure + profocc + clerocc + servocc, data = wage1)
wage.lm9 <- lm(wage ~ educ + exper + tenure + female, data = wage1)
wage.lm11 <- lm(wage ~ educ + exper + tenure + female*married, data = wage1)
wage.lm12 <- lm(log(wage) ~ educ + exper + tenure + female*married, data = wage1)
wage.lm13 <- lm(wage ~ educ + exper + tenure + female*educ, data = wage1)
crime1 <- wooldridge::crime1
crime.lm1 <- lm((narr86 > 0) ~ pcnv + avgsen + tottime + ptime86 + qemp86, data = crime1)
gradethis::gradethis_setup()
knitr::opts_chunk$set(echo = FALSE)
```

## Dummy (Binary) Independent Variables

Qualitative Information (Categorical/Factor Variables)

  - Gender, Race, Industry, Occupation
  - Create separate dummy variables for each category. 
  - Avoid the *dummy variable trap*!
  - R automatically does this for you. 
  
![](https://media.giphy.com/media/lk0TFUdop2JTW/giphy.gif)

## Single Dummy Independent Variable

Estimate the wage differential for women, controlling for education, experience, and job tenure. Call this $wage.lm9$ and run a $summary$ of the results. 

```{r gender, exercise = TRUE}

```

```{r gender-hint}
wage.lm9 <- lm(..., data = ...)
summary(...)
```

```{r gender-solution}
wage.lm9 <- lm(wage ~ educ + exper + tenure + female, data = wage1)
summary(wage.lm9)
```

```{r gender-check}
grade_code()
```

## Dummy Variables for Multiple Categories

- Use the $factor()$ function to create a factor, $occ$, as a variable in $wage1$ from the occupational dummy variables in $wage1$ ($profocc$, $clerocc$, and $servocc$) that takes integer values from one to four. Use the $labels$ option to assign the labels $Manufacturing$, $Professional$, $Clerical$, and $Services$ to the values. 
- Replicate the regression, wage.lm7 using the new factor variable you created. Call this $wage.lm10$.  
- Summarize both regressions using a text stargazer() output.

```{r occupation, exercise = TRUE}

```

```{r occupation-hint}
# The variables profocc, clerocc, and servocc take values on {0,1}.
# To make them to each take sequential integer values, you can just multiply them by a constant.
# Get them to take values {0,1} for profocc, {0,2} for clerocc, and {0,3 for servocc}. 
# Add one to scale it from 1 through 4 (more natural for categories) instead of 0 through 3. 
wage1$occ <- factor(..., labels = c("Manufacturing", "Professional", "Clerical", "Services"))
wage.lm7  <- lm(..., data = ...)
wage.lm10 <- lm(..., data = ...)
stargazer(..., type = "text")
```

```{r occupation-solution}
wage1$occ <- factor(1 + wage1$profocc + 2*wage1$clerocc + 3*wage1$servocc, levels = c("Manufacturing", "Professional", "Clerical", "Services"))
wage.lm7 <- lm(wage ~ educ + exper + I(exper^2) + tenure + profocc + clerocc + servocc, data = wage1)
wage.lm10 <- lm(wage ~ educ + exper + I(exper^2) + tenure + occ, data = wage1)
stargazer(wage.lm7, wage.lm10, type = "text")
```

```{r occupation-check}
grade_code("They're the same!")
```

## Interactions among Dummy Variables

- Estimate the wage effect of marriage differ across gender by interacting $female$ with $married$, controlling for education, experience, and job tenure. Call this $wage.lm11$ 
- Estimate the effect of marriage on *log wages* across gender with the same controls. Call this $wage.lm12$ 
- Print a text $stargazer()$ of your results. 

```{r marriage, exercise = TRUE}

```

```{r marriage-hint}
wage.lm11 <- lm(..., data = ...)
wage.lm12 <- lm(..., data = ...)
stargazer(..., type = 'text')
```

```{r marriage-solution}
wage.lm11 <- lm(wage ~ educ + exper + tenure + female*married, data = wage1)
wage.lm12 <- lm(log(wage) ~ educ + exper + tenure + female*married, data = wage1)
stargazer(wage.lm11, wage.lm12, type = 'text')
```

```{r marriage-check}
grade_code()
```

## Different Slopes

What if we think the *effect of education* (a continuous variable in our model) differs for women? 

- Estimate the effect of gender on wages *and the returns to education*, controlling for experience and job tenure. Call this $wage.lm13$ and run a $summary()$ of your result. 

```{r genderEducation, exercise = TRUE}

```

```{r genderEducation-hint}
wage.lm13 <- lm(..., data = ...)
summary(...)
```

```{r genderEducation-solution}
wage.lm13 <- lm(wage ~ educ + exper + tenure + female*educ, data = wage1)
summary(wage.lm13)
```

```{r genderEducation-check}
grade_code()
```

## Binary Dependent Variables: The Linear Probability Model

- Estimate the effect of prior convictions ($pcnv$) on whether a person was arrested in 1986, controlling for the average sentence of prior convictions ($avesen$), total time spent in prison prior to 1986 ($tottime$), total number of months spent in prison in 1986 ($ptime86$), and total number of quarters officially employed in 1986 ($qemp86$). 
- Call this $crime.lm1$ and run a $summary()$ of your result.

```{r arrests, exercise = TRUE}

```

```{r arrests-hint}
# We don't have the variable we need in the data, but we don't need to create it!
# The formula argument understands logical functions. 
# Since we have the number of times a person was arrested, narr86, we can operate a logical condition within our lm formula argument for narr86 > 0 and specify this directly as the dependent variable. 
crime.lm1 <- lm(..., data = ...)
summary(...)
```

```{r arrests-solution}
crime.lm1 <- lm((narr86 > 0) ~ pcnv + avgsen + tottime + ptime86 + qemp86, data = crime1)
summary(crime.lm1)
```

```{r arrests-check}
grade_code()
```

What does the coefficient on prior convictions represent? 
