---
title: "Chapter 13"
subtitle: "Pooled Cross Sections and Difference-in-Differences Estimation"
description: >
  This tutorial introduces the OLS model for panel data and difference-in-differences.
author: 
  name: "Created by [Jim Bang](http://www.github.com/bangecon)"
  email: BangJamesT@sau.edu
  affiliation: "[St. Ambrose University](http://www.sau.edu)"
date: "`r format(Sys.Date(), '%d %b %Y')`"
output: 
  learnr::tutorial:
    progressive: true
    allow_skip: true
runtime: shiny_prerendered
always_allow_html: yes
editor_options: 
  markdown: 
    wrap: 72
---

<style type="text/css">
h1{font-size: 24pt}
h2{font-size: 20pt}
h3{font-size: 18pt}
h4,h5,h6{font-size: 16pt}
body{font-size: 16pt}
</style>

```{r setup, include=FALSE}
library(learnr)
library(gradethis)
tutorial_options(exercise.eval = TRUE, exercise.reveal_solution = TRUE)
gradethis::gradethis_setup()
knitr::opts_chunk$set(echo = FALSE, fig.height= 9, fig.width=8)
```

## Pooled Cross Sections vs. Panel (Longitudinal) Data

Pooled Cross Sections (Quasi-Panel) - Data obtained by conducting a cross-sectional survey multiple times over different time periods with independent random samples each period.

-   Current Population Survey
-   World Values Survey (Individual Responses)
-   General Social Survey

Panel (Longitudinal) Data - Data obtained by tracking the same (randomly-sampled) individuals over multiple time periods.

-   Panel Study of Income Dynamics
-   World Development Indicators
-   World Values Survey (Country Averages)
-   Various Business Sources at:

<https://maryannfeldman.web.unc.edu/data-sources/longitudinal-databases/>

## Balanced vs. Unbalanced Panels

Balanced Panel: Same entities tracked over the exact same time periods. No missing observations.

Unbalanced Panel: Same entities tracked over time, but some entities may have different numbers of time periods than others, or gaps between some years. 

## Pooled Cross Sections Exercise

The $cps78_85$ dataset contains independent random samples of wages and worker characteristics for Current Population Surveys taken in 1978 and 1985.

The individuals in the 1978 sample are not the same as the ones in the 1985 sample.

Using a single line of code, regress, name ($cps78_85.lm1$), and summarize the linear model of log wages on education, gender, experience and its square, union status, and the interactions of year = 1985 with education and gender.

```{r cps, exercise = TRUE}

```

```{r cps-hint}
# You can nest your lm command inside the summary command, with or without naming and storing the result. 
summary(cps78_85.lm1 <- lm(..., ...))
```

```{r cps-solution}
summary(cps78_85.lm1 <- lm(log(wage) ~ y85*(educ+female) + exper + I(exper^2) + union, data=cps78_85))
```

```{r cps-check}
grade_code()
```

## Simple Difference-in-Difference (DiD) Regression

Did a garbage incinerator lower house prices, or was it built where house prices were already lower?

To investigate, estimate the following using the $kielmc$ data: 

  1. $ln(RealPrice_i^{1978}) = \beta_0^{1978} + \beta_1^{1978}Incinerator_i^{1978} + u_i^{1978}$; $incinerator.lm78$
  2. $ln(RealPrice_i^{1981}) = \beta_0^{1981} + \beta_1^{1978}Incinerator_i^{1981} + u_i^{1981}$; $incinerator.lm78$
  3. $ln(RealPrice_{it}) = \beta_0 + \delta_0Year^{1981} + \beta_1Incinerator_i + \delta_1Year^{1981}*Incinerator_i + u_{it}$; $incinerator.DiD1$

```{r incinerator, exercise = TRUE}

```

```{r incinerator-hint}

```

```{r incinerator-solution}
# Separate regressions for 1978 and 1981: report coeeficients only
incinerator.lm78 <- lm(log(rprice) ~ nearinc, data=kielmc, subset=(year==1978))
incinerator.lm81 <- lm(log(rprice) ~ nearinc, data=kielmc, subset=(year==1981))
# Joint regression including an interaction term 
incinerator.DiD1 <- lm(log(rprice) ~ nearinc*factor(year), data=kielmc)
stargazer(incinerator.lm78, incinerator.lm81, incinerator.DiD1, type = 'text')
```

```{r incinerator-check}
grade_code()
```

## DiD with Controls

Estimate the effect of the incinerator using the DiD method, controlling for age of the house, $age$, and its square; the log of the distance of the house to the interstate, $intst$; the log of the lot's square footage, $land$; the log of the house's square footage, $area$; the total number of rooms in the house, $rooms$; and the number of bathrooms, $baths$. 

```{r DiD, exercise = TRUE}

```

```{r DiD-hint}

```

```{r DiD-solution}
incinerator.DiD2 <- lm(log(rprice) ~ nearinc*factor(year) + age + I(age^2) + log(intst) + log(land) + log(area) + rooms + baths, data = kielmc)
stargazer(incinerator.DiD1, incinerator.DiD2, type  = 'text')
```

```{r DiD-check}
grade_code()
```
