---
title: "Chapter 4-1 Practice"
output: learnr::tutorial
runtime: shiny_prerendered
---

```{r setup, include=FALSE}
library(learnr)
library(gradethis)
library(wooldridge)
crime2 <- wooldridge::crime2
gradethis_setup()
tutorial_options(exercise.timelimit = 60)
knitr::opts_chunk$set(echo = FALSE)
```

## Population Density and Crime 

1. Estimate the following model using the `crime2` dataset.
$$ln(CrimeRate) = \beta_0 + \beta_1ln(LawExpendPC) + \beta_2ln(IncomePC) + \beta_3ln(PopulationDensity) + \beta_4West + \beta5Northeast + \beta_6South + u$$
2. Construct a 99% confidence interval for the expenditure elasticity of crime rate using `confint`.

```{r crime, exercise=TRUE}

```

```{r crime-solution, exercise.reveal_solution = FALSE}
crime.lm1 <- lm(log(crmrte) ~ log(lawexpc) + log(pcinc) + log(popden) + west + nrtheast + south, data = crime2)
confint(crime.lm1, "log(lawexpc)", level = 0.99)
```

```{r crime-check}
grade_this({if (!inherits(.result, c("matrix"))) {
    fail("Your class of your answer should be a `matrix` object.")}
  if (round(.result$coef[1],4) != 0.1535) { 
    fail("The coefficients on education should be about 0.15") }
  pass()
})
```
