---
title: "Chapter 2-2 Practice"
output: learnr::tutorial
runtime: shiny_prerendered
---

```{r setup, include=FALSE}
library(learnr)
library(gradethis)
ceosal1 <- wooldridge::ceosal1
meap93 <- wooldridge::meap93
gradethis_setup()
tutorial_options(exercise.timelimit = 60)
knitr::opts_chunk$set(echo = FALSE)
```

## CEO Salaries and Sales

Use the `ceosal1` data to regress the *log* of CEO salaries on the *log* of sales. Name the results `ceosal.lm`

Summarize the results and plot the data with the regression line. 

```{r lmPractice, exercise=TRUE, exercise.reveal_solution = FALSE}

```

```{r lmPractice-solution}
ceosal.lm <- lm(log(salary) ~ log(sales), data = ceosal1)
summary(ceosal.lm)
plot(log(ceosal1$sales),
     log(ceosal1$salary),
     xlab = "ln(Sales)",
     ylab = "ln(Salary)")
abline(ceosal.lm)
```

```{r lmPractice-check}
grade_code()
```

